# Copyright 2026 Ubuntu
# See LICENSE file for licensing details.

rule_files:
  - ../../src/prometheus_alert_rules/alerts.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: 'ovs_build_info{}'
        values: '1x10'
    alert_rule_test:
      - eval_time: 6m
        alertname: MicroOVNExporterDown
        exp_alerts: []

  - interval: 1m
    input_series: []
    alert_rule_test:
      - eval_time: 6m
        alertname: MicroOVNExporterDown
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              summary: OVN metrics endpoint is unavailable
              description: No OVN exporter metrics are being scraped for 5 minutes.

  - interval: 1m
    input_series:
      - series: 'ovn_controller_southbound_database_connected{}'
        values: '0x10'
    alert_rule_test:
      - eval_time: 6m
        alertname: OvnControllerSouthboundDisconnected
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              summary: ovn-controller disconnected from SB DB
              description: ovn-controller is not connected to the OVN southbound database.

  - interval: 1m
    input_series:
      - series: 'ovn_northd_status{instance="a"}'
        values: '1x10'
      - series: 'ovn_northd_status{instance="b"}'
        values: '1x10'
    alert_rule_test:
      - eval_time: 6m
        alertname: OvnNorthdNoActiveInstance
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              summary: invalid ovn-northd active/standby state
              description: Expected exactly one active ovn-northd instance, got 2.

  - interval: 1m
    input_series:
      - series: 'ovn_northd_nb_connection_status{}'
        values: '0x10'
      - series: 'ovn_northd_sb_connection_status{}'
        values: '0x10'
    alert_rule_test:
      - eval_time: 6m
        alertname: OvnNorthdNBConnectionDown
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              summary: ovn-northd disconnected from NB DB
              description: ovn-northd is not connected to the OVN northbound database.
      - eval_time: 6m
        alertname: OvnNorthdSBConnectionDown
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              summary: ovn-northd disconnected from SB DB
              description: ovn-northd is not connected to the OVN southbound database.

  - interval: 1m
    input_series:
      - series: 'ovn_db_cluster_server_role{db_name="OVN_Northbound", server_role="leader", server_id="a"}'
        values: '0x10'
    alert_rule_test:
      - eval_time: 6m
        alertname: OvnNBDBNoLeader
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              summary: OVN northbound DB has no leader
              description: No OVN northbound database leader is currently reported.

  - interval: 1m
    input_series:
      - series: 'ovn_db_cluster_server_role{db_name="OVN_Northbound", server_role="leader", server_id="a"}'
        values: '1x10'
      - series: 'ovn_db_cluster_server_role{db_name="OVN_Northbound", server_role="leader", server_id="b"}'
        values: '1x10'
    alert_rule_test:
      - eval_time: 6m
        alertname: OvnNBDBMultipleLeaders
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              summary: OVN northbound DB has multiple leaders
              description: More than one OVN northbound database leader is currently reported.

  - interval: 1m
    input_series:
      - series: 'ovn_db_cluster_server_role{db_name="OVN_Southbound", server_role="leader", server_id="a"}'
        values: '0x10'
    alert_rule_test:
      - eval_time: 6m
        alertname: OvnSBDBNoLeader
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              summary: OVN southbound DB has no leader
              description: No OVN southbound database leader is currently reported.

  - interval: 1m
    input_series:
      - series: 'ovn_db_cluster_server_role{db_name="OVN_Southbound", server_role="leader", server_id="a"}'
        values: '1x10'
      - series: 'ovn_db_cluster_server_role{db_name="OVN_Southbound", server_role="leader", server_id="b"}'
        values: '1x10'
    alert_rule_test:
      - eval_time: 6m
        alertname: OvnSBDBMultipleLeaders
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              summary: OVN southbound DB has multiple leaders
              description: More than one OVN southbound database leader is currently reported.

  - interval: 1m
    input_series:
      - series: 'ovn_db_cluster_log_not_applied{db_name="OVN_Northbound"}'
        values: '1x10'
      - series: 'ovn_db_cluster_log_not_committed{db_name="OVN_Northbound"}'
        values: '0x10'
    alert_rule_test:
      - eval_time: 11m
        alertname: OvnDBClusterLogBacklog
        exp_alerts:
          - exp_labels:
              severity: warning
            exp_annotations:
              summary: OVN DB cluster has unapplied/uncommitted log entries
              description: OVN DB cluster replication is lagging and has pending log entries.

  - interval: 1m
    input_series:
      - series: 'ovn_northd_txn_error{}'
        values: '0 1 2 3 4 5 6 7 8 9 10 11 12 13 14 15 16 17 18 19 20'
    alert_rule_test:
      - eval_time: 20m
        alertname: OvnNorthdTransactionErrors
        exp_alerts:
          - exp_labels:
              severity: warning
            exp_annotations:
              summary: ovn-northd transaction errors detected
              description: ovn-northd reported transaction errors in the last 10 minutes.

  - interval: 1m
    input_series:
      - series: 'ovn_controller_txn_error{}'
        values: '0 1 2 3 4 5 6 7 8 9 10 11 12 13 14 15 16 17 18 19 20'
    alert_rule_test:
      - eval_time: 20m
        alertname: OvnControllerTransactionErrors
        exp_alerts:
          - exp_labels:
              severity: warning
            exp_annotations:
              summary: ovn-controller transaction errors detected
              description: ovn-controller reported transaction errors in the last 10 minutes.

  - interval: 1m
    input_series:
      - series: 'ovn_controller_txn_incomplete{}'
        values: '0 1 2 3 4 5 6 7 8 9 10 11 12 13 14 15 16 17 18 19 20'
    alert_rule_test:
      - eval_time: 20m
        alertname: OvnControllerTransactionIncomplete
        exp_alerts:
          - exp_labels:
              severity: warning
            exp_annotations:
              summary: ovn-controller incomplete transactions detected
              description: ovn-controller reported incomplete transactions in the last 10 minutes.

  - interval: 1m
    input_series:
      - series: 'ovs_vswitchd_process_open_fds{}'
        values: '950x10'
      - series: 'ovs_vswitchd_process_max_fds{}'
        values: '1000x10'
    alert_rule_test:
      - eval_time: 11m
        alertname: OvsVswitchdFileDescriptorsHigh
        exp_alerts:
          - exp_labels:
              severity: warning
            exp_annotations:
              summary: ovs-vswitchd file descriptor usage is high
              description: ovs-vswitchd is using more than 90% of available file descriptors.
